---
title: 'Hypothesis Testing'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Hypothesis Testing

```yaml
type: VideoExercise
key: 58b1b669b4
xp: 50
```

`@projector_key`
de75eb5d8613c1adaedd3593b32f4583
